# Esercizi per casa

## Addizione
#### Fai la somma dei seguenti numeri:

<!-- grade_id: somma, points: 3, grade: true, solution: true -->
2+2=?

## Sottrazione
#### Fai la differenza tra i seguenti numeri:

<!-- grade_id: sottrazione, points: 4, grade: true, solution: true -->
4-3=?

## Moltiplicazione - autograded - answer
#### scrivi una funzione che fa il prodotto

<!-- grade_id: moltiplicazione, points: 0, grade: false, solution: true -->
```python
def moltiplicazione(a,b):
    """Questa funzione deve restituire il prodotto tra due numeri"""
    
    
    ###begin solution
    if a==0 or b==0:
        raise ValueError("neither a nor b must be 0")
    return a*b
    ###end solution
```

## Divisione - autograded - test
#### Quanto fa 6/3?

<!-- grade_id: divisione, points: 4, grade: true, solution: false -->
```python
from nose.tools import assert_equal

def divisione(a,b):
    ###begin solution
    return a/b
    ###end solution
    
assert_equal(divisione(6,3),2)
```

```python

```